# Implement Standard Library Functions
---
This task focuses on building OCaml functions to mirror OCaml's standard library. Don't look at those ones! 

The first function to implement is `map` - this applies a supplied function `f` to every element of a given list. Make sure your solution can cope with big lists!

<!-- grade_id: map-function, points: 0, grade: false, solution: true -->
```OCaml
(** [map f lst] should return a list of equal length with [f] applied to every element of [lst]*)
let map f lst = 
    (* BEGIN SOLUTION *)
    let rec aux acc = function  
      | [] -> List.rev acc 
      | x::xs -> aux (f x :: acc) xs 
    in 
      aux [] lst 
    (* END SOLUTION *)
```

Output:
```
val map : ('a -> 'b) -> 'a list -> 'b list = <fun>
```

<!-- grade_id: map-function-check, points: 5, grade: true, solution: false -->
```OCaml
assert ((map (fun x -> x + 1) [0; 1; 2; 3; 4]) = [1; 2; 3; 4; 5]);
assert ((map string_of_int [0; 1; 2]) = ["0"; "1"; "2"]); 
(* BEGIN HIDDEN TESTS *)

assert (let lst = List.init 1_000_000 (fun x -> x) in (map (fun x -> x) lst = lst)); 
(* END HIDDEN TESTS *)
```

Output:
```
- : unit = ()
```

```OCaml

```