Make sure to fill in the meta-information before you begin. You should replace `(* Your Solution Here *)` with your solutions and make sure the assrtions all pass. 

```OCaml
let name = ""
let suid = ""
```

# Implement Standard Library Functions
---
This task focuses on building OCaml functions to mirror OCaml's standard library. Don't look at those ones! 

The first function to implement is `map` - this applies a supplied function `f` to every element of a given list. Make sure your solution can cope with big lists!

<!-- grade_id: map-function, points: 0, grade: false, solution: true -->
```OCaml
(** [map f lst] should return a list of equal length with [f] applied to every element of [lst]*)
let map f lst = 
    (* Your Solution Here *)
     assert false
```

<!-- grade_id: map-function-check, points: 5, grade: true, solution: false -->
```OCaml
assert ((map (fun x -> x + 1) [0; 1; 2; 3; 4]) = [1; 2; 3; 4; 5]);
assert ((map string_of_int [0; 1; 2]) = ["0"; "1"; "2"]); 
```

```OCaml

```